```python
version = "REPLACE_PACKAGE_VERSION"
```

---
# Assignment 2 Part 1: Vector Space Model and Relevance Feedback (50 pts)

In this assignment, we will improve our simple text retrieval system built in Assignment 1 by leveraging the Vector Space Model and relevance feedback. 

```python
# Configure nltk

import nltk

nltk_data_path = "assets/nltk_data"
if nltk_data_path not in nltk.data.path:
    nltk.data.path.append(nltk_data_path)
    
```

```python
# ignores warnings, should be on for submitting to the autograder but you may comment it out if needed while working.
import warnings
warnings.simplefilter("ignore")
```

## Question 1: Retrieve and rank documents with TF-IDF weighting scheme (10 pts)

Modify your `retrieve_n_rank_docs` function from Assignment 1 so that it now uses the TF-IDF weights to retrieve and rank documents as shown in the lecture slides. Specifically, for each *term* in a given query, we accumulate the TF-IDF weight for each document that contains the term. We apply **maximum frequency normalisation** to raw term frequencies to obtain the TF and follow the definition of IDF given in the lecture. After repeating this procedure for all terms, we rank the documents found in descending order of their total TF-IDF weights **rounded to three decimal places**. Again, if two documents tie, the document with a *lower* document number should be ranked higher. Refer to the lecture slides titled "Ranking Documents: Example" in the TF-IDF slide deck for more details. 

Your modified function should accept the same set of arguments as in Assignment 1 and likewise output a `dict` that contains an ordered `list` of retrieved documents for each query, similar to:

```
{
    'q1'  : ['d51', 'd486', ..., 'd876'], 
    'q2'  : ['d12', 'd51', ..., 'd486'],
    ...,
    'q225': ['d1188', 'd1380', ..., 'd173']
}
```

⚠️ Please remember to **round the total TF-IDF weights to three decimal places** before ranking the documents. ⚠️

**This function should return a `dict` of length `len(queries)`, where each key is a `str` representing a `q_id` and each value is a `list` that contains the `doc_id` of the retrieved documents. `max_docs` determines the maximum number of documents that can be retrieved for each query.**

<!-- grade_id: cell-79211162e032488b, points: 0, grade: false, solution: true -->
```python
import math

def retrieve_n_rank_docs(inverted_index, queries, max_docs=-1):
    """
    Retrieve documents in order of relevance from an inverted index based on some queries.
    """
    
        # Calculate maximum frequency for each document
    max_freq = {}
    for token, doc_data in inverted_index.items():
        for doc_id, freq_tuple in doc_data.items():
            freq = freq_tuple[0]  # Assuming the frequency is the first element in the tuple
            if doc_id not in max_freq:
                max_freq[doc_id] = freq
            else:
                max_freq[doc_id] = max(max_freq[doc_id], freq)

    # Total count of documents
    N = len({doc_key for doc_data in inverted_index.values() for doc_key in doc_data})
    
    ret_docs ={}
    for q_id, terms in queries.items():
        doc_scores = {}
        for term in terms:
            if term in inverted_index:
                k = len(inverted_index[term])
                idf = 1 + math.log(N/k)
                for doc_id, freq_tuple in inverted_index[term].items():
                    freq = freq_tuple[0]
                    tf = 0.5 + (0.5 * freq / max_freq[doc_id])
                    if doc_id not in doc_scores:
                        doc_scores[doc_id] = 0
                    doc_scores[doc_id] = doc_scores[doc_id] + (tf * idf)
        sorted_docs = sorted(doc_scores.items(), key=lambda x: (-round(x[1], 3), int(x[0][1:])))
        # Select top documents based on max_docs
        top_docs = [doc_id for doc_id, _ in sorted_docs][:max_docs if max_docs != -1 else None]
        ret_docs[q_id] = top_docs

    return ret_docs

```

```python
# We import the reference implementations of some utility functions from Week 1 to be used in the autograder tests
# They should be used ONLY in the autograder tests, NOT in your solutions
# Ignore any ModuleNotFoundError message in the autograder feedback - it does not affect your grade
from utils import load_cranfield_docs, load_cranfield_queries, build_inverted_index
```

<!-- grade_id: cell-5b892e216b90d5e7, points: 0, grade: true, solution: false -->
```python
# The definitions of the three utility functions are hidden here - not for grading

```

<!-- grade_id: cell-66bf33f44f8b83bf, points: 10, grade: true, solution: false -->
```python
# Autograder tests

min_df = 10 # min_df won't change in the hidden tests
stu_docs = load_cranfield_docs()
stu_inv_index = build_inverted_index(stu_docs, min_df=min_df)
stu_queries = load_cranfield_queries()

max_docs = 10 # max_docs may vary in the hidden tests
stu_ret_docs = retrieve_n_rank_docs(stu_inv_index, stu_queries, max_docs=max_docs)

# Some sanity checks
assert isinstance(stu_ret_docs, dict), "Q1: Your function should return a dictionary. "
assert len(stu_ret_docs) == len(stu_queries), "Q1: Your dictionary should have the same length as there are queries. "

for q_id in stu_ret_docs:
    
    assert q_id in stu_queries, f"Q1: When max_docs = {max_docs}, '{q_id}' in your dictionary is not a valid q_id. "
    
    assert len(stu_ret_docs[q_id]) <= max_docs, f"Q1: When max_docs = {max_docs}, your # retrieved docs ({len(stu_ret_docs[q_id])}) for {q_id} is bigger than max_docs. "

# Some hidden tests

del stu_inv_index, stu_queries, stu_ret_docs, min_df, max_docs
```

## Question 2: Rocchio Relevance Feedback (40 pts)

Now comes the really fun stuff. Let's try to simulate an interactive Rocchio relevance feedback loop. The two main characters in the story are:

* an inquisitive `User` who never ceases issuing queries; and
* a diligent `RetrievalSystem` that always attends to the queries with extraordinary patience.

You will play the role of `RetrievalSystem`, represented by the `RetrievalSystem` class. An instance of the `User` class can be created as follows.  

```python
from utils import User
test_user = User(random_state=0) # The User in the autograder tests uses a different seed
```

### Question 2a: Perform Latent Semantic Indexing (LSI) (10 pts)

As the first step of constructing a `RetrievalSystem` class, let's first complete the `__init__` method. It accepts a collection of documents, `docs`, as returned by the `load_cranfield_docs` function, and an argument `num_concepts` indicating the number of concepts/topics to keep for LSI. In addition to the argument `min_df` with its usual meaning, the method also accepts three other arguments, `alpha`, `beta` and `gamma`, that correspond to the $\alpha$, $\beta$ and $\gamma$ parameters in the Rocchio query update formula. 

A document-term matrix with TF-IDF weights, `doc_term_mat`, has been created for you. Your task at this step is to perform LSI on the document-term matrix and obtain the LSI-transformed vectors in the concept space for each document. The LSI-transformed vectors should be stored in the variable `self.doc_vecs` as a 2-D `np.ndarray` of the shape `(len(docs), num_concepts)`. You may (and probably need to) create additional class attributes to hold objects that may be shared with other class methods. Please use `random_state=42` whenever it is required. 


**This class method should fill the class attribute `self.doc_vecs` with a 2-D `np.ndarray` of the shape `(len(docs), num_concepts)`.**

<!-- grade_id: cell-38d49c4872fe96ab, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict
import numpy as np
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.decomposition import TruncatedSVD


class RetrievalSystem:
    def __init__(self, docs, num_concepts, min_df=1, alpha=1.0, beta=0.75, gamma=0.15):
        self.alpha, self.beta, self.gamma = alpha, beta, gamma
        self.vec = TfidfVectorizer(tokenizer=str.split, min_df=min_df)
        doc_term_mat = self.vec.fit_transform([" ".join(docs[doc_id]) for doc_id in docs])
        self.q_vecs = {}
        svd = TruncatedSVD(n_components=num_concepts, random_state=42)
        self.doc_vecs = svd.fit_transform(doc_term_mat)
        
        

```

<!-- grade_id: cell-ce481913f67a4bf0, points: 10, grade: true, solution: false -->
```python
# Autograder tests

num_concepts, min_df = 100, 10 # These won't change in the hidden tests

stu_ret_sys = RetrievalSystem(stu_docs, num_concepts, min_df)

# Some sanity checks
assert isinstance(stu_ret_sys.doc_vecs, np.ndarray), "Q2a: The doc_vecs attribute of your RetrievalSystem should be a np.ndarray. "
assert stu_ret_sys.doc_vecs.shape == (len(stu_docs), num_concepts), "Q2a: The doc_vecs attribute of your RetrievalSystem has an incorrect shape. "
assert np.issubdtype(stu_ret_sys.doc_vecs.dtype, np.floating), "Q2a: The doc_vecs attribute of your RetrievalSystem should have a float dtype. "


# Some hidden tests

del stu_ret_sys, num_concepts, min_df
```

### Question 2b: Retrieve and rank documents in the concept space (10 pts)

Next, let's complete the method `retrieve_n_rank_docs` for retrieving and ranking documents in the concept space generated by your LSI earlier. The method accepts the same arguments `queries` and `max_docs` as you have seen in other retrieval functions. The only novelty is that now we keep the LSI-transformed vectors for all the documents instead of an inverted index. 

The argument `queries` is a `dict` of queries as returned by the `load_cranfield_queries` function. However, it may contain an arbitrary subset of the Cranfield queries rather than necessarily all the Cranfield queries. In fact, you should not assume any knowledge about the possible queries included in `queries`, because what queries you will receive is at the discretion of the inquisitive `User`. Upon receiving `queries`, you should turn each **NEW** query into an LSI-tranformed vector, following what you have done with the documents; for queries that you have received before, fetch their vectors from the class attribute `self.q_vecs`. The reason why we need to treat new and old queries differently is that later we will perform Rocchio updates on some query vectors so that their final vector representations will differ from the LSI representations they started with. All query vectors start as an LSI vector but may end up differently depending on how often they are involved in the feedback loop. 

Once you have the correct vector representation for each query, retrieve documents in descending order of the cosine similarity between their vector representations. The maximum number of documents to retrieve for each query is again governed by the argument `max_docs`. As before, your method should finally output a `dict` containing the documents retrieved for each query, similar to:

```
{
    'q217': ['d983', 'd554', ..., 'd623'],
    'q99' : ['d716', 'd67', ..., 'd164'],
    ...
}
```

It is fine if you see a `RuntimeWarning`. That's because two document vectors are zero vectors. 


⚠️ **Always remember, the document numbers start at 1.** Namely, the first document is `d1`. ⚠️

**This method should return a `dict` of length `len(queries)`, where each key is a `str` representing a `q_id` and each value is a `list` that contains the `doc_id` of the retrieved documents in order. `max_docs` determines the maximum number of documents to retrieve for each query.**

<!-- grade_id: cell-3f5bce7fe99fdc6b, points: 0, grade: false, solution: true -->
```python
class RetrievalSystem:
    
    def __init__(self, docs, num_concepts, min_df=1, alpha=1.0, beta=0.75, gamma=0.15):
        self.alpha, self.beta, self.gamma = alpha, beta, gamma
        self.vec = TfidfVectorizer(tokenizer=str.split, min_df=min_df)
        doc_term_mat = self.vec.fit_transform([" ".join(docs[doc_id]) for doc_id in docs])
        self.q_vecs = {}
        self.svd = TruncatedSVD(n_components=num_concepts, random_state=42)
        self.doc_vecs = self.svd.fit_transform(doc_term_mat)


    def retrieve_n_rank_docs(self, queries, max_docs=-1):
        """
        Retrieve and rank documents in the latent semantic (concept) space
        """
        
        ret_docs = {}
        for q_id, q_terms in queries.items():
            if q_id not in self.q_vecs:
                q_vec = self.vec.transform([" ".join(q_terms)])
                self.q_vecs[q_id] = self.svd.transform(q_vec)
            cos_sim = []
            for doc_vec in self.doc_vecs:
                sim = np.dot(self.q_vecs[q_id], doc_vec.T)/(np.linalg.norm(self.q_vecs[q_id])*np.linalg.norm(doc_vec))
                cos_sim.append(sim)
            cos_sim = np.array(cos_sim).flatten()
            ranked_docs = np.argsort(-cos_sim)[:max_docs if max_docs != -1 else None]
            ret_docs[q_id] = [f'd{idx + 1}' for idx in ranked_docs]
                
        
        return ret_docs

```

<!-- grade_id: cell-94b2a07cd98ffa25, points: 10, grade: true, solution: false -->
```python
# Autograder tests
import random
random.seed(0)

num_concepts, min_df = 100, 10 # These won't change in the hidden tests
max_docs = 10 # This may vary in the hidden tests

queries = load_cranfield_queries()
stu_queries = dict(random.sample(queries.items(), 100)) # Test on a random sample

stu_ret_sys = RetrievalSystem(stu_docs, num_concepts, min_df)
stu_ret_docs = stu_ret_sys.retrieve_n_rank_docs(stu_queries, max_docs)

# Some sanity checks
assert isinstance(stu_ret_docs, dict), "Q2b: Your method should return a dictionary. "
assert len(stu_ret_docs) == len(stu_queries), "Q2b: Your dictionary should have the same length as there are queries. "

for q_id in stu_ret_docs:
    
    assert q_id in stu_queries, f"Q2b: When max_docs = {max_docs}, '{q_id}' in your dictionary is not a valid q_id. "
    
    assert len(stu_ret_docs[q_id]) <= max_docs, f"Q2b: When max_docs = {max_docs}, your # retrieved docs ({len(stu_ret_docs[q_id])}) for {q_id} is bigger than max_docs. "


# Some hidden tests

del stu_ret_sys, num_concepts, min_df
del max_docs, queries, stu_queries, stu_ret_docs
```

### Question 2c: Perform Rocchio query updates (20 pts)

Finally it's time to engage with the inquisitive `User`. Complete the method `gather_feedback` for inviting a `user` to your Rocchio relevance feedback loop. Your interaction with the `user` proceeds in four steps. 

* **Step 1**: You request the `user` to issue queries by invoking the `user`'s `issue_queries` method. The `user` then returns to you a `dict` of queries represented by `queries`, and a number `max_docs` indicating the maximum number of documents the `user` wants to retrieve for each query. 


* **Step 2**: You retrieve the required number of documents based on `queries` as you did in the last question. 


* **Step 3**: You invite the `user` to give feedback to the documents retrieved, `ret_docs`, by invoking the `user`'s `give_feedback` method with `ret_docs` passed along. Normally in a relevance feedback loop, the `user` would mark each document retrieved as either relevant or irrelevant; however, the `user` you are engaging with prefers giving feedback in the form of **precisions at each rank**. 
    
    Specifically, suppose your `ret_docs` is as follows:

    ```
    {
        'q167': ['d553', 'd1100', 'd1096', 'd1279', 'd1099'], 
        'q54' : ['d123', 'd554', 'd623', 'd398', 'd102'], 
        'q197': ['d768', 'd884', 'd883', 'd909', 'd882']
    }
    ```
    from which we can deduce that the `user` must have issued three queries and for each query at most five documents should be retrieved. If you pass your `ret_docs` to the `user` for feedback, you would get back:
    
    ```
    {
        'q167': [0.0, 0.0, 0.0, 0.0, 0.0], 
        'q54' : [0.0, 0.0, 0.0, 0.0, 0.0], 
        'q197': [0.0, 0.5, 0.3333333333333333, 0.25, 0.2]
    }
    ```
    where each number indicates the precision at that rank. For example, the second number `0.5` from `q197` indicates that the precision at rank 2 is `0.5` for query `q197`. Similarly, the precision at rank 4 is `0.25` for query `q197`. For queries `q167` and `q54`, the precisions at top 5 ranks are, unfortunately, zero. The precisions at each rank are returned back to you in the variable `pre_at_n`, along with another variable `avg_ndcg` which indicates your average NDCG for all queries. 
    
Code for the first three steps has been provided to you. Your task is to complete the code for

* **Step 4**: You perform a Rocchio update on all the queries the `user` issued to you based on the feedback you received. As a result, the query vectors in `self.q_vecs` should be updated accordingly. In other words, the "modified query vector" for each query under consideration should be stored in the corresponding entry in `self.q_vecs`. 

⚠️ **Always remember, the document numbers start at 1.** Namely, the first document is `d1`. ⚠️

**This method should return a copy of `self.q_vecs` for grading.**

<!-- grade_id: cell-a8150e5145aa8eed, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict
import numpy as np
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.decomposition import TruncatedSVD


class RetrievalSystem:
    
    def __init__(self, docs, num_concepts, min_df=1, alpha=1.0, beta=0.75, gamma=0.15):
        self.alpha, self.beta, self.gamma = alpha, beta, gamma
        self.vec = TfidfVectorizer(tokenizer=str.split, min_df=min_df)
        doc_term_mat = self.vec.fit_transform([" ".join(docs[doc_id]) for doc_id in docs])
        self.q_vecs = {}
        self.svd = TruncatedSVD(n_components=num_concepts, random_state=42)
        self.doc_vecs = self.svd.fit_transform(doc_term_mat)


    def retrieve_n_rank_docs(self, queries, max_docs=-1):
        """
        Retrieve and rank documents in the latent semantic (concept) space
        """
        
        ret_docs = {}
        for q_id, q_terms in queries.items():
            if q_id not in self.q_vecs:
                q_vec = self.vec.transform([" ".join(q_terms)])
                self.q_vecs[q_id] = self.svd.transform(q_vec)
            cos_sim = []
            for doc_vec in self.doc_vecs:
                sim = np.dot(self.q_vecs[q_id], doc_vec.T)/(np.linalg.norm(self.q_vecs[q_id])*np.linalg.norm(doc_vec))
                cos_sim.append(sim)
            cos_sim = np.array(cos_sim).flatten()
            ranked_docs = np.argsort(-cos_sim)[:max_docs if max_docs != -1 else None]
            ret_docs[q_id] = [f'd{idx + 1}' for idx in ranked_docs]
    
        return ret_docs
    

    def gather_feedback(self, user):
        """
        This function models the interactive relevance feedback loop.
        """
        queries, max_docs = user.issue_queries()
        ret_docs = self.retrieve_n_rank_docs(queries, max_docs=max_docs)
        pre_at_n, avg_ndcg = user.give_feedback(ret_docs)
        
        for q_id in queries:
            relevant_docs_indices = []
            non_relevant_docs_indices = []
            prev_precision_product = 0
            rank = 1
            for doc_id, precision in zip(ret_docs[q_id], pre_at_n[q_id]):
                doc_index = int(doc_id[1:]) - 1
                precision_product = rank * precision
                if precision_product > prev_precision_product:
                    relevant_docs_indices.append(doc_index)
                else:
                    non_relevant_docs_indices.append(doc_index)
                prev_precision_product = precision_product
                rank += 1
            relevant_vectors = self.doc_vecs[relevant_docs_indices] if relevant_docs_indices else np.zeros(self.doc_vecs.shape[1])
            non_relevant_vectors = self.doc_vecs[non_relevant_docs_indices] if non_relevant_docs_indices else np.zeros(self.doc_vecs.shape[1])
            query_vec = self.q_vecs[q_id]
            rel_component = np.mean(relevant_vectors, axis=0) if relevant_docs_indices else np.zeros(self.doc_vecs.shape[1])
            non_rel_component = np.mean(non_relevant_vectors, axis=0) if non_relevant_docs_indices else np.zeros(self.doc_vecs.shape[1])
            self.q_vecs[q_id] = self.alpha * query_vec + self.beta * rel_component - self.gamma * non_rel_component
                

        return self.q_vecs.copy()
```

<!-- grade_id: cell-70974b1dcf8b7f3d, points: 0, grade: true, solution: false -->
```python
# The definition of the User class is hidden here - not for grading

```

<!-- grade_id: cell-97527c3c993bc9a6, points: 20, grade: true, solution: false -->
```python
# Autograder tests

num_concepts, min_df = 100, 10 # These won't change in the hidden tests

stu_ret_sys = RetrievalSystem(stu_docs, num_concepts, min_df)


# Some hidden tests

del stu_ret_sys, num_concepts, min_df
```

```python

```